# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd
import numpy as np
import re
```

```python
cdc = pd.read_csv(r"D:/Introduction-to-Data-Science-in-Python/assignment2/assets/NISPUF17.csv", index_col = 0)
cdc.head(3)
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT  PROVWT_D     RDDWT_D  STRATUM  YEAR  AGECPOXR  \
1   128521     12852     2       NaN  235.916956     1031  2017       NaN   
2    10741      1074     2       NaN  957.353840     1068  2017       NaN   
3   220011     22001     2       NaN  189.611299     1050  2017       NaN   

   HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  XVRCTY6  \
1         2       1  ...                        NaN      NaN      NaN   
2         2       1  ...                        NaN      NaN      NaN   
3         2       3  ...                        NaN      NaN      NaN   

   XVRCTY7  XVRCTY8  XVRCTY9  INS_STAT2_I  INS_BREAK_I  
1      NaN      NaN      NaN          NaN          NaN  
2      NaN      NaN      NaN          NaN          NaN  
3      NaN      NaN      NaN          NaN          NaN  

[3 rows x 453 columns]
```

```python
cdc.shape
```

Output:
```
(28465, 454)
```

```python
[col_name for col_name in cdc.columns if "ED" in col_name]
```

Output:
```
['EDUC1']
```

```python
cdc.EDUC1.min(), cdc.EDUC1.max()
```

Output:
```
(1, 4)
```

```python
sum(cdc.EDUC1.isna())
```

Output:
```
0
```

### ***EDUC1***
***1 - <12 years       
2 - 12 years      
3 - >12 years, not a college graduate        
4 - College graduate***

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
        result = dict.fromkeys(["less than high school", "high school", "more than high school but not college", "college"])
        
        result["less than high school"] = len(cdc.query("EDUC1 == 1")) / cdc.shape[0]
        result["high school"] = len(cdc.query("EDUC1 == 2")) / cdc.shape[0]
        result["more than high school but not college"] = len(cdc.query("EDUC1 == 3")) / cdc.shape[0]
        result["college"] = len(cdc.query("EDUC1 == 4")) / cdc.shape[0]
        
        return result
```

```python
proportion_of_education()
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
cdc.CBF_01.unique()
```

Output:
```
array([ 1,  2, 99, 77], dtype=int64)
```

```python
cdc.P_NUMFLU.unique()
```

Output:
```
array([nan,  3.,  0.,  2.,  1.,  4.,  5.,  6.])
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
# CBF_01 – child ever fed breast milk
# 1 - Yes
# 2 - No
# 77 - Don’t know
# 99 - Missing

# P_NUMFLU – total number of seasonal influenza doses

def average_influenza_doses():
    return cdc.query("CBF_01 == 1").loc[:, "P_NUMFLU"].mean(), cdc.query("CBF_01 == 2").loc[:, "P_NUMFLU"].mean()
    
```

```python
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
cdc.SEX.unique()
```

Output:
```
array([1, 2], dtype=int64)
```

```python
cdc.HAD_CPOX.unique()
```

Output:
```
array([ 2,  1, 77, 99], dtype=int64)
```

```python
cdc.P_NUMVRC.unique()
```

Output:
```
array([nan,  1.,  0.,  2.,  3.])
```

```python
cdc.query("(HAD_CPOX == 1) & (P_NUMVRC >= 1) & (SEX == 1)").P_NUMVRC
```

Output:
```
462      1.0
688      1.0
864      1.0
1480     1.0
1907     1.0
        ... 
27046    1.0
27721    1.0
28107    1.0
28108    1.0
28405    1.0
Name: P_NUMVRC, Length: 68, dtype: float64
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
# SEX - sex of the child
# 1 - Male
# 2 - Female

# HAD_CPOX – did child ever have chicken pox
# 1 - Yes
# 2 - No
# 77 - Don’t know
# 99 - Missing

# P_NUMVRC – total number of varicella doses


def chickenpox_by_sex():
    stats = dict.fromkeys(["male", "female"])
    
    stats["male"] = cdc.query("(HAD_CPOX == 1) & (P_NUMVRC >= 1) & (SEX == 1)").shape[0] / \
                                    cdc.query("(HAD_CPOX == 2) & (P_NUMVRC >= 1) & (SEX == 1)").shape[0]
    stats["female"] = cdc.query("(HAD_CPOX == 1) & (P_NUMVRC >= 1) & (SEX == 2)").shape[0] / \
                                    cdc.query("(HAD_CPOX == 2) & (P_NUMVRC >= 1) & (SEX == 2)").shape[0]
    return stats
```

```python
chickenpox_by_sex()
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df = pd.DataFrame({"had_chickenpox_column": cdc.HAD_CPOX,
                   "num_chickenpox_vaccine_column": cdc.P_NUMVRC}).dropna(axis = 0)

    # here is some stub code to actually run the correlation
    corr, pval = stats.pearsonr(df["had_chickenpox_column"], df["num_chickenpox_vaccine_column"])
    
    print(pval)
    # just return the correlation
    return corr

    # YOUR CODE HERE
    # raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
corr_chickenpox()
```

Output:
```
0.6669755758983259
0.0034753171515294265
```

```python
df = pd.DataFrame({"had_chickenpox_column": cdc.HAD_CPOX,
                   "num_chickenpox_vaccine_column": cdc.P_NUMVRC}).dropna(axis = 0)
df.reset_index(drop = True, inplace = True)
```

```python
df.had_chickenpox_column.value_counts(77)
```

Output:
```
2     0.985652
1     0.011283
77    0.003065
Name: had_chickenpox_column, dtype: float64
```

```python
df.isnull().apply(sum, axis = 0)
```

Output:
```
had_chickenpox_column            0
num_chickenpox_vaccine_column    0
dtype: int64
```

```python
df.drop(np.where(df.had_chickenpox_column == 77)[0], inplace = True)
```

```python

```

```python

```