# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    # raise NotImplementedError()
    return re.findall("[A-Z][a-z]+[^\s\,]", simple_string)
```

```python
names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
    # raise NotImplementedError()
    return [student.split(":")[0] for student in re.findall("\w+\s\w+:\sB", grades)]
```

```python
grades(), len(grades())
```

Output:
```
(['Bell Kassulke',
  'Simon Loidl',
  'Elias Jovanovic',
  'Hakim Botros',
  'Emilie Lorentsen',
  'Jake Wood',
  'Fatemeh Akhtar',
  'Kim Weston',
  'Yasmin Dar',
  'Viswamitra Upandhye',
  'Killian Kaufman',
  'Elwood Page',
  'Elodie Booker',
  'Adnan Chen',
  'Hank Spinka',
  'Hannah Bayer'],
 16)
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

```python
with open("./assets/logdata.txt", "r") as file:
    logs_ = file.read().split("\n")
logs_.pop()
```

Output:
```
''
```

```python
# all ip addresses have just three periods!
{log.split(" ")[0].count(".") for log in logs_}
```

Output:
```
{3}
```

```python
from collections import Counter
```

```python
# there are 478 hyphens inplace of user names!
Counter([log.split(" ")[2] for log in logs_]).get("-")
```

Output:
```
478
```

```python
# 4 numbers separated by three periods, exclusing the trailing space
ip_regex = "\d+\.\d+\.\d+\.\d+"

# starting with a hyphen, followed by a space and the user name (empty user names are handled)
# trim out the hyphen & space
user_regex = "-\s[a-zA-Z0-9\-]*[^\s]"

#              dd    /mmm    /yyyy :hh   :mm    :ss     -0700
time_regex = "\d{2}\/[a-zA-Z]{3}\/\d{4}\:\d{2}\:\d{2}\:\d{2}\s-\d{4}"   # "\d{2}\/\w{3}/\/d{4}\:\d{2}\:\d{2}\:\d{2}\s\-\d{4}"

# exclude the starting & trailing double quotes, and accomodate the rest, including the forward slashes
request_regex = "[^\"][A-Z]+\s\/[\w\/\s\.\-]+[^\"]"
```

```python
len([re.findall(ip_regex, log)[0] for log in logs_])
```

Output:
```
979
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
def logs():
    
    # a list to hold the serialized logs (as dictionaries)
    serialized_data = list()
    
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read().split("\n")
    
    logdata.pop() # remove the last space
    
    for line_no, log in enumerate(logdata):
        # print(f"{line_no}: ✓", sep = "", end = ",  ")
        log_dict = {
            "host": re.findall(ip_regex, log)[0],
            "user_name": re.findall(user_regex, log)[0].replace("- ", ""),
            "time": re.findall(time_regex, log)[0],
            "request": re.findall(request_regex, log)[0]
                   }
        serialized_data.append(log_dict)
    return serialized_data
```

```python
len(logs())
```

Output:
```
979
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
logs()[0] == one_item
```

Output:
```
True
```

```python
one_item in logs()
```

Output:
```
True
```

```python
import re

def logs():
    serialized_data = list()
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read().split("\n")
    logdata.pop()
    
    ip_regex = "\d+\.\d+\.\d+\.\d+"
    user_regex = "-\s[a-zA-Z0-9\-]*[^\s]"
    time_regex = "\d{2}\/[a-zA-Z]{3}\/\d{4}\:\d{2}\:\d{2}\:\d{2}\s-\d{4}" 
    request_regex = "[^\"][A-Z]+\s\/[\w\/\s\.\-]+[^\"]"
    
    for line_no, log in enumerate(logdata):
        # print(f"{line_no}: ✓", sep = "", end = ",  ")
        log_dict = {
            "host": re.findall(ip_regex, log)[0],
            "user_name": re.findall(user_regex, log)[0].replace("- ", ""),
            "time": re.findall(time_regex, log)[0],
            "request": re.findall(request_regex, log)[0]
                   }
        serialized_data.append(log_dict)
    return serialized_data
```

```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
logs_ = logs()
```

```python
from collections import Counter
```

```python
{record.get("host").count(".") for record in logs_}
```

Output:
```
{3}
```

```python
len(logs_)
```

Output:
```
979
```

```python
# copy paste from a learner

def logs_cpaste():
    with open("./assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    host_username = re.findall("[\d]*.[\d]*.[\d].*[\d]\s-\s[\w]*-*", logdata)
    time = re.findall("\[(.*)\]", logdata)
    request = re.findall("[A-Z][A-Z]+\s*\.*\n*.+[A-Z]*\.[0-9]", logdata)
    
    res = []
    i = 0
    for item in host_username:
        host = re.split("[ - ]", item)[0]
        username = re.split("[ - ]", item)[2]
        res.append({
            "host": host,
            "user_name": username,
            "time": time[i],
            "request": request[i]
        })
        i += 1
    return res
```

```python
my_logs = logs()
cpy_logs = logs_cpaste()

{log == log_ for (log, log_) in zip(my_logs, cpy_logs)}
```

Output:
```
{True}
```

```python

```